# Working with multiple data sets

There are two data files that we'll be working with for this week's assignment.  They are described below.  Load those data files in with Pandas and then work to answering each of the questions below.

## npidata.csv

This file is basic information about every healthcare provider in the US.  It has one row for each NPI (National Provider Identifier).  It contains information such as the provider's name and address.


## cmsYYYY.csv

These are files about what kinds of procedures and patients providers in the US are serving under CMS programs, Medicare and Medicaid.  Each contains various statistics about providers over the course of a year.  There are three of these, for the years 2014, 2015, and 2016.  These files, however, don't contain information about the provider such as where the provider is located.

## Our Goals

For this assignment, we're going to want to compute some statistics based on the data in the **cms** files, but aggregate that data based on information in the **npidata** file.  As in last week's assignment, you'll need to store your answers in a variable called `answer` at the end of each step.

```python
import pandas as pd
```

## Part 1

In this first step, we'll need to merge together all of the **cms** files into a single dataframe.  Be careful that these files might not be identical, so you'll have to look a little bit to figure out how to merge them.

As you are merging them, make sure that you retain information about which file (i.e. which year) the data came from.  Call that new columns `year`.

In your `answer` variable, provide a complete data frame that contains all of the rows and columns from the **cms** files, plus an additional column to store the year/file that particular row came from.

The assertion tests will give you a good idea as to if you're merging the files correctly.

<!-- grade_id: 01-solution, points: 0, grade: false, solution: true -->
```python
cms_file_2014 = pd.read_csv("../../data/cms2014.csv")
cms_file_2015 = pd.read_csv("../../data/cms2015.csv")
cms_file_2016 = pd.read_csv("../../data/cms2016.csv")
npi_data = pd.read_csv("../../data/npidata.csv")
```

Output:
```
/opt/tljh/user/lib/python3.6/site-packages/IPython/core/interactiveshell.py:3051: DtypeWarning: Columns (18,33) have mixed types.Specify dtype option on import or set low_memory=False.
  interactivity=interactivity, compiler=compiler, result=result)
```

```python
answer =None
cms_file_2014['year'] = 2014
cms_file_2015['year'] = 2015
cms_file_2016['year'] = 2016
```

```python
cms_file_2014.columns = cms_file_2014.columns.str.strip().str.lower()
cms_file_2015.columns = cms_file_2015.columns.str.strip().str.lower()
cms_file_2016.columns = cms_file_2016.columns.str.strip().str.lower()
answer = pd.concat([cms_file_2014, cms_file_2015, cms_file_2016], ignore_index=True)
```

<!-- grade_id: 01-tests, points: 5, grade: true, solution: false -->
```python
assert(answer.shape == (193862, 60))
assert(list(answer['year'].unique()) == [2014,2015,2016])
assert(set(answer.columns.str.lower()) == set(['year', 'nbr', 'npi', 'provider_type',
       'medicare_participation_indicator', 'number_of_hcpcs', 'total_services',
       'total_unique_benes', 'total_submitted_chrg_amt',
       'total_medicare_allowed_amt', 'total_medicare_payment_amt',
       'total_medicare_stnd_amt', 'drug_suppress_indicator',
       'number_of_drug_hcpcs', 'total_drug_services',
       'total_drug_unique_benes', 'total_drug_submitted_chrg_amt',
       'total_drug_medicare_allowed_amt', 'total_drug_medicare_payment_amt',
       'total_drug_medicare_stnd_amt', 'med_suppress_indicator',
       'number_of_med_hcpcs', 'total_med_services', 'total_med_unique_benes',
       'total_med_submitted_chrg_amt', 'total_med_medicare_allowed_amt',
       'total_med_medicare_payment_amt', 'total_med_medicare_stnd_amt',
       'beneficiary_average_age', 'beneficiary_age_less_65_count',
       'beneficiary_age_65_74_count', 'beneficiary_age_75_84_count',
       'beneficiary_age_greater_84_count', 'beneficiary_female_count',
       'beneficiary_male_count', 'beneficiary_race_white_count',
       'beneficiary_race_black_count', 'beneficiary_race_api_count',
       'beneficiary_race_hispanic_count', 'beneficiary_race_natind_count',
       'beneficiary_race_other_count', 'beneficiary_nondual_count',
       'beneficiary_dual_count', 'beneficiary_cc_afib_percent',
       'beneficiary_cc_alzrdsd_percent', 'beneficiary_cc_asthma_percent',
       'beneficiary_cc_cancer_percent', 'beneficiary_cc_chf_percent',
       'beneficiary_cc_ckd_percent', 'beneficiary_cc_copd_percent',
       'beneficiary_cc_depr_percent', 'beneficiary_cc_diab_percent',
       'beneficiary_cc_hyperl_percent', 'beneficiary_cc_hypert_percent',
       'beneficiary_cc_ihd_percent', 'beneficiary_cc_ost_percent',
       'beneficiary_cc_raoa_percent', 'beneficiary_cc_schiot_percent',
       'beneficiary_cc_strk_percent', 'beneficiary_average_risk_score']))
```

```python

```

## Part 2

In this next part, we're going to join the **cms** data with the provider information in the **/data/npidata.csv** file.  In this join, we don't want to lose any records from the **cms** files, even if no matching provider exists in the **npidata** file.  However, we don't care about any providers from the **npidata** file that don't have records in the **cms** files.  Those providers can be ignored.

Join the data files together to create one unified data frame called `answer`.  This dataframe should have all the columns from both **cms** and **npidata** files, joined together using the `npi` column.  

Note that `npi` is unique in the **npidata** file.

<!-- grade_id: 02-solution, points: 0, grade: false, solution: true -->
```python
import pandas as pd


answer = pd.merge(answer, npi_data, how='left', left_on='npi', right_on='NPI')



```

<!-- grade_id: 02-tests, points: 5, grade: true, solution: false -->
```python
import numpy
assert(answer.shape == (193862, 103))
assert(list(answer['Provider Business Mailing Address State Name'].unique()) == ['IL','MO',numpy.nan,'WY'])
assert(list(answer.groupby('Provider Business Mailing Address State Name').NPI.count()) == [111520, 53366, 4805])
```

## Part 3

If you did everything right above, you'll notice that grouping by a column with NaN in it will cause some rows to disappear from the aggregation test.  So, let's create a new column called `'State'` that has the same value as whatever is in the `'Provider Business Mailing Address State Name'` column or a value of `'XX'` if there is no state information.

Set `answer` to be your final data frame with the new `'State'` column added.

<!-- grade_id: 03-solution, points: 0, grade: false, solution: true -->
```python
answer['State'] = answer['Provider Business Mailing Address State Name'].fillna('XX')
```

```python
answer
```

Output:
```
nbr         npi                    provider_type  \
0             1  1003000134                        Pathology   
1            19  1003001249               Physical Therapist   
2            21  1003001322  Mass Immunization Roster Biller   
3            37  1003002486                     Chiropractic   
4            73  1003005430            Clinical Psychologist   
...         ...         ...                              ...   
193857  1053890  1992995542                    Ophthalmology   
193858  1053892  1992995591          Obstetrics & Gynecology   
193859  1053900  1992996219                        Neurology   
193860  1053901  1992996284                Internal Medicine   
193861  1053906  1992996458                   Sleep Medicine   

       medicare_participation_indicator  number_of_hcpcs  total_services  \
0                                     Y               13          8643.0   
1                                     Y                4          1221.0   
2              
```

<!-- grade_id: 03-tests, points: 5, grade: true, solution: false -->
```python
assert(list(answer.groupby('State').npi.count()) == [111520, 53366, 4805, 24171])
assert(answer.shape == (193862, 104))
```

## Part 4

Next, let's summarize the data by year and by State.  Create a pivot table that contains one row for each state and one column for each year.  Within the pivot table, put a sum of total services as the values.

Assign `answer` to be that resulting pivot table.  In the tests, I'm going to plot a bar chart of your pivot table.

<!-- grade_id: 04-solution, points: 0, grade: false, solution: true -->
```python
%matplotlib inline
final_dt = answer.groupby(['year','State'], as_index=False)['total_services'].sum()
final_dt = final_dt.pivot(index='State', columns='year', values='total_services')

```

```python
answer = final
```

<!-- grade_id: 04-tests, points: 5, grade: true, solution: false -->
```python
assert(answer.shape == (4,3))
assert(answer.sum().sum() == 519185664.6999999)
assert(answer[2016].sum() == 176596933.80000004)
assert(answer.loc['WY'].sum() == 10892707.4)
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-12-db10226393b0> in <module>
----> 1 assert(answer.shape == (4,3))
      2 assert(answer.sum().sum() == 519185664.6999999)
      3 assert(answer[2016].sum() == 176596933.80000004)
      4 assert(answer.loc['WY'].sum() == 10892707.4)

AssertionError:
```

```python
%matplotlib inline
answer.plot.bar()
```

```python

```



---



## Submitting Your Work

In order to submit your work, you'll need to use the `git` command line program to **add** your homework file (this file) to your local repository, **commit** your changes to your local repository, and then **push** those changes up to github.com.  From there, I'll be able to **pull** the changes down and do my grading.  I'll provide some feedback, **commit** and **push** my comments back to you.  Next week, I'll show you how to **pull** down my comments.

To run through everything one last time and submit your work:
1. Use the `Kernel` -> `Restart Kernel and Run All Cells` menu option to run everything from top to bottom and stop here.
2. Follow the instruction on the prompt below to either ssave and submit your work, or continue working.

If anything fails along the way with this submission part of the process, let me know.  I'll help you troubleshoort.

---

```python
a=input('''
Are you ready to submit your work?
1. Click the Save icon (or do Ctrl-S / Cmd-S)
2. Type "yes" or "no" below
3. Press Enter

''')

if a=='yes':
    !git pull
    !git add week14_assignment.ipynb
    !git commit -a -m "Submitting the week 14 programming assignment"
    !git push
else:
    print('''
    
OK. We can wait.
''')
```

```python

```